# Problem Set 04 (20 pts)
## Due Monday May 4, 2026
### Topics:
* Mortgages and Plotting
* Numpy
* Probability and Plotting

### Question 1 (1 pt, manually graded): List Comprehension

Using list comprehension, list all of the numbers from 1-105 that are divisible by 7.

<!-- grade_id: cell-697c6d10d66d9b79, points: 1, grade: true, solution: false -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 2 (1pt, manually graded): Strings & List Comprehension

Use list comprehension to count the number of spaces in the string `"This is a test string"`

<!-- grade_id: cell-186d1fe869ee4770, points: 1, grade: true, solution: false -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 3 (2pts, manually graded): Strings & List Comprehension

Use the `split()` function on a string to see what it does when you don't pass any parameters. 

Then, use this and list comprehension to make a list that only keeps the words that are less than or equal to four letters in the string 

`"Poincaré's homology sphere is a closed 3- manifold with the same homology as the 3-sphere but with a fundamental group which is non-trivial."` 

The output should be

`['is', 'a', '3-', 'the', 'as', 'the', 'but', 'a', 'is']`

<!-- grade_id: cell-89279d098e0ec19b, points: 2, grade: true, solution: false -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 4 (4 pts, manually graded): Mortgages

Produce a plot that shows how much is owed on a $250,000 loan with

1) a fixed 5% loan for 30 years, and 

2) a two rate mortgage that has a 3.5% rate for 5 years and then 7% for the remaining 25 years

Your figure should look like the following, with the x-axis label saying "Number of Months" the y-axis label saying "Balance Remaining" and the legend describing the two curves.

![Mortgage Comparison](https://github.com/yiboxu20/ProgrammingDS/blob/main/Notebooks/figs/Mortgage_Comparison.png?raw=true)

I've included the code from lecture below. Make sure to import any libraries you may need for plotting.

```python
def findPayment(loan,r,m):
    """Assumes: loan and r are floats, m an int
     Returns the monthly payment for a mortgage of size
     loan at a monthly rate of r for m months"""
    return loan*(r*(1+r)**m)/((1+r)**m -1)

class Mortgage(object):
    """Abstract class for building different kinds of mortgages"""
    
    def __init__(self, loan, annRate, months):
        """Assumes: loan and annRate are floats, months an int.
        Creates a new mortgage of size loan, duration months, and
        annual rate annRate"""
        self.loan = loan
        self.rate = annRate/12.0
        self.months = months
        self.paid = [0.0]
        self.owed = [loan]
        self.payment = findPayment(loan, self.rate, months)
        self.legend = None #description of mortgage
        
    def makePayment(self):
        """Make a payment"""
        self.paid.append(self.payment)
        reduction = self.payment - self.owed[-1]*self.rate
        self.owed.append(self.owed[-1] - reduction) #remember that [-1] returns the last element of a list
    
    def getTotalPaid(self):
        """Return the total amount paid so far"""
        return sum(self.paid)
    
    def __str__(self):
        return self.legend
    
class Fixed(Mortgage):
    def __init__(self, loan, r, months):
        Mortgage.__init__(self, loan, r, months)
        self.legend = 'Fixed, ' + str(r*100) + '%'

class FixedWithPts(Mortgage):
    
    def __init__(self, loan, r, months, pts):
        Mortgage.__init__(self, loan, r, months)
        self.pts = pts
        self.paid = [loan*(pts/100.0)]
        self.legend = 'Fixed, ' + str(r*100) + '%, '\
                    + str(pts) + ' points'
class TwoRate(Mortgage):
    def __init__(self, loan, r, months, teaserRate, teaserMonths):
        Mortgage.__init__(self, loan, teaserRate, months)
        self.teaserMonths = teaserMonths
        self.teaserRate = teaserRate
        self.nextRate = r/12.0
        self.legend = str(teaserRate*100)\
                + '% for ' + str(self.teaserMonths)\
                + ' months, then ' + str(round(r*100,2)) + '%'
    
    def makePayment(self):
        if len(self.paid) == self.teaserMonths + 1:
            self.rate = self.nextRate
            self.payment = findPayment(self.owed[-1], self.rate, self.months - self.teaserMonths)
        
        Mortgage.makePayment(self)
```

<!-- grade_id: cell-bae0e83e94a37782, points: 4, grade: true, solution: false -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 5 (4 pts, manually graded): Mortgages

This is a math question and an opportunity to practice your $\LaTeX$. A derivation of the formula is required for full credit. 

*N.B. Finding a closed form formula is hard, so don't feel ashamed to find a computational solution, which is worth 3 points*

Assuming a fixed rate mortage of rate $r$ with a loan period of $N$ months.

**In what month $k$ have you paid off half the original (principal) loan amount, but in month $k-1$ you still owe more than half of the principal?** 

The ideal answer should be a function of $N$ and $r$.

<!-- grade_id: cell-11e024d4c17f0c42, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

### Question 6 (2 pts, autograded): NumPy

Complete the function below so that an integer $n$ is taken in as input and returns a NumPy array of the odd numbers between 0 and n (inclusive).

```python
import numpy as np
def beOdd(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
beOdd(10)
```

```python
np.array([1, 3, 5, 7, 9])
```

```python

```

<!-- grade_id: cell-b0f585eb851a0950, points: 2, grade: true, solution: false -->
```python
#TEST CASES
np.testing.assert_array_equal(beOdd(10),np.array([1, 3, 5, 7, 9]))
np.testing.assert_array_equal(beOdd(30),np.array([ 1,  3,  5,  7,  9, 11, 13, 15, 17, 19, 21, 23, 25, 27, 29]))
```

### Question 7 (2 pts, autograded): NumPy

Complete the function below so that an integer $n$ is taken as input and returns a Numpy array of the numbers between 0 and n (inclusive), but where every entry that is congruent to 1 modulo 3 is replaced with -42.

```python
def replace1mod3(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-761fb1be341f9dc6, points: 2, grade: true, solution: false -->
```python
np.testing.assert_array_equal(replace1mod3(10),np.array([0,-42,2,3,-42,5,6,-42,8,9,-42]))
np.testing.assert_array_equal(replace1mod3(15),np.array([0,-42,2,3,-42,5,6,-42,8,9,-42,11,12,-42,14,15]))
```

### Question 8 (4pts, manually graded): Rolling a Die Twice
#### Part A
Write a function that calculates the probability of rolling exactly two 3's in k rolls of a fair six sided die. 


<!-- grade_id: cell-03a0c8a87e2db1a6, points: 2, grade: true, solution: false -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Part B
Use the function from Q8, Part A to plot the probability as k varies from 2 to 100.

<!-- grade_id: cell-cb52863b7497db0d, points: 2, grade: true, solution: false -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```